# CVAI - Übung: Farben

### Detektion mit Farben

In unserem SkinApp Projekt befassen wir uns mit der Erkennung von Hautkrankheiten, zum Beispiel auf Händen. Als erstes muss dabei die Handfläche vom Hintergrund getrennt werden. Damit dies einfacher geht, haben wir eine Photobox entwickelt, die die Hände vor grünem Hintergrund aufnimmt. Sie sollen nun untersuchen, ob es damit tatsächlich einfacher ist, die Handflächen oder Handrückseiten zu finden.





```python
# OpenCV needs to be included first

import cv2
import numpy as np

# for displaying images in jupyter

import matplotlib as mpl
from matplotlib import pyplot as plt

%matplotlib inline
# resolution of images in figure
mpl.rcParams['figure.dpi']= 200
# figure size in inches
mpl.rcParams["figure.figsize"] = [10,5]

path = '/exchange/cvai/images/'
```

### Aufgabe 1: Handsegmentierung

Als erstes sollen sie versuchen, die Handflächen mittels der Farbe der Haut zu detektieren. Dazu ist es möglicherweise hilfreich, das Bild in einen anderen Farbraum zu transferieren. Für Hautsegmentierung könnte sich zum Beispiel der YCrCb- oder der HSV-Farbraum eignen. Stellen Sie das Bild in diesen Farbräumen dar und suchen Sie dann, entsprechende Schwellwerte um eine binäre Maske der Hand zu erstellen. Für Schwellwerte stehen in OpenCV eigentlich die Funktionen `cv.threshold()` und `cv2.inRange()` zur Verfügung, für einfache Fälle ist oft auch ein Vergleich direkt in numpy möglich.

```python
image_grey_back = cv2.imread (path + 'hand_grey_small.png')
image_grey_back = cv2.cvtColor (image_grey_back, cv2.COLOR_BGR2RGB)

# simple threshold in RGB

min = np.array ([150, 0, 0])
max = np.array ([255, 140, 150])

skin_image = cv2.inRange (image_grey_back, min, max)

plt.subplot (1, 2, 1); plt.imshow (image_grey_back)
plt.subplot (1, 2, 2); plt.imshow (skin_image, cmap = 'gray')
```

Konvertieren sie nun das Bild in ein anderes Farbsystem und versuchen sie, dort einen besseren Schwellwert zu erhalten.

<!-- grade_id: cell-17b506beb75e8020, points: 0, grade: false, solution: true -->
```python
image_hsv = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3ec33c22a1b32f7a, points: 1, grade: true, solution: false -->
```python
assert image_hsv is not None
#plt.imshow(image_hsv)
```

### Aufgabe 2: Hintergrundsegmentierung

Im folgenden wird ein Bild mit grünem Hintergrund geladen. Führen Sie darauf die gleiche Segmentierung wie beim oberen Bild durch und stellen Sie diese dar.

```python
image_green_back = cv2.imread (path + 'hand_green_small.png')
image_green_back = cv2.cvtColor (image_green_back, cv2.COLOR_BGR2RGB)
plt.imshow (image_green_back)
```

Segmentieren Sie nun statt der Haut den grünen Hintergrund. Können Sie damit eine bessere Unterscheidung zwischen Hand und Hintergrund erhalten? Stellen sie die beiden Bilder nebeneinander dar.

<!-- grade_id: cell-4ffa178be8476e1f, points: 0, grade: false, solution: true -->
```python
hsv = None

# YOUR CODE HERE
raise NotImplementedError()

#plt.subplot (1, 2, 1); plt.imshow (..., cmap = 'gray')
#plt.subplot (1, 2, 2); plt.imshow (..., cmap = 'gray')
```

<!-- grade_id: cell-83f05625f4e8a43c, points: 1, grade: true, solution: false -->
```python
assert hsv is not None
#plt.imshow(hsv)
```

### Aufgabe 3: Bluescreen

Blue Screening wird in der Filmindustrie verwendet, um Regionen in einem Bild (zum Beispiel einen blauen Hintergrund) durch ein anderes Bild zu ersetzen.

Entwickeln sie ein Program, das zwei Bilder mit Hilfe dieser Technik überlagert. Gehen Sie dazu wie folgt vor:

* Berechnen sie für jedes Pixel die Transparenz α aufgrund des Farbwertes des ersten Bildes (I1)

* Berechnen sie die Pixel des Ausgabebildes O, indem sie und das zweite Bild (I2) mit Hilfe des α-Wertes übereinander legen. Eine mögliche Funktion dazu ist 

$$ O = \alpha  \cdot I_1 + (1 − \alpha) \cdot I_2 $$

* Dies wird als Alpha Blendung bezeichnet.

Ein Beispielbild mit blauem Hintergrund, sowie ein Beispiel Hintergrund werden im anschliessenden Code geladen. Allerdings haben die Bilder noch nicht die gleiche Grösse und müssen also zuerst noch skaliert werden. Sie können aber auch andere Bilder aus dem Internet verwenden.

```python
foreground = cv2.imread (path + '2011-04-29_KG_RLB_RED_LEADER-003C_small.jpg')
foreground = cv2.cvtColor (foreground, cv2.COLOR_BGR2RGB)

background = cv2.imread (path + 'landscape-moon.jpg')
background = cv2.cvtColor (background, cv2.COLOR_BGR2RGB)

plt.subplot (1, 2, 1); plt.imshow (foreground)
plt.subplot (1, 2, 2); plt.imshow (background)
print (foreground.shape)
print (background.shape)
```

<!-- grade_id: cell-ce6b91741c247146, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4b13835b31324cd9, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19cd092f596657a2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec13899a761383c3, points: 1, grade: true, solution: false -->
```python
# This should be 0
print(foreground.shape[0] - background.shape[0])
print(foreground.shape[1] - background.shape[1])

# These should be equal
print(foreground.shape)
print(background.shape)

np.testing.assert_equal(foreground.shape, background.shape)
```

```python

```

```python

```